```python
# Initialize Otter
import otter
grader = otter.Notebook("bondwindowdressing.ipynb")
```

# Quantitative Research Lab: Discovering Trade Ideas

This lab assignment aims to help guide you through the process of analyzing data to discover potential trading ideas in financial markets

# Importing Libraries

```python
import yfinance as yf
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import warnings
warnings.filterwarnings('ignore')
```

# Data Collection

To start, we're going to collect the data we need to analyze the effect. 

We'll use yfinance, a popular python library for accessing the yahoo finance API for asset price data

```python
# BEGIN SOLUTION NO PROMPT
import yfinance as yf

# Define the ticker symbol for TLT
ticker_symbol = "TLT"

start_date = "2002-12-31"  
end_date = "2023-01-01"    

# Use yfinance to download the data
tlt_data = yf.download("TLT", start=start_date, end=end_date)
tltraw = pd.DataFrame(tlt_data)
tltraw

# END SOLUTION
""" # BEGIN PROMPT
import yfinance as yf

# Define the ticker symbol as TLT
ticker_symbol = "TLT"

start_date = "2002-12-31"  
end_date = "2023-01-01"    

# fill in the correct parameters for the download function
# yf.download (asset_symbol_to_download, start= ..., end=...)
tlt_data = yf.download(..., start=..., end=...)

# turn the data into a dataframe
tltraw = pd.DataFrame()
tltraw

""" # END PROMPT
```

```python
grader.check("q1")
```

let's plot the data, to see how it looks!

```python
tltraw['Close'].plot()
```

```python
tltraw['Adj Close'].plot()
```

```python
# adjusted close is basically just the close price, adjusted for stock splits and dividends
# the adjusted close price is what you would actually be able to trade, so that's what we look at
```

### Log vs Simple Returns

```python
def simplereturn(x0,x1):
    return ...
def logreturn(x0,x1):
    return ...
def logtosimple(r):
    return ...
def simpletolog(r):
    return ...

tltdf = tltraw[['Adj Close']]

tltdf['logreturn'] = logreturn()
tltdf
tltdf['logreturn'].cumsum().plot()
```

```python
grader.check("q2")
```

# Exploring Effect

since we're looking at an effect that happens during the beginning and end of a month, we'll want to start by grouping our data based on the day of the month

```python


tltdf['dayofmonth'] = ... # use the index of the dataframe, and the .day accessor
dom_mean_logret = tltdf.groupby(...).function()['column'] # input the correct column to group on, column to view, and aggregation function

```

```python
grader.check("q3")
```

now that we've grouped the data successfully, we'll look at the average logreturn by day, to see if anything interesting pops out

# what do you expect to see in the data? 

```python
plt.bar(dom_mean_logret.index, dom_mean_logret)
```

# what do you notice in the data? 

# coming up with a basic trading rule

```python


# we'll first want to group the data by months, as we're studying a month by month effect

# you can pass lists into groupby!

tltdf_months = tltdf.groupby([tltdf.index., tltdf.index.])

# now, we want to create a function that we can apply to each group
# test out what one of the groups looks like, see what september 2022 looks like

sep2022 = tltdf_months.get_group((...,...))

# try to get the sum of the log returns of the first and last 5 days
# hint: use the pandas iloc accessor

last_5_long_922 = sep2022
first_5_short_922 = sep2022

```

```python
grader.check("q4")
```

calculate what would happen if we were long during the last five months, and short during the first five months

```python

# generalize the above into a function that we can run on every group (test it on september 2022: you should get the same number as before)
def calculate_returns(group, n_days):
    # when we're long, last n days
    last_5_sum = 
    # when we're short, first n days
    first_5_sum = 
    # combine
    return last_5_sum + first_5_sum
    
calculate_returns(sep2022, 5)

# now let's apply it to every month in the data
monthly_returns = tltdf_months.apply(calculate_returns, 5)
monthly_returns

```

```python
grader.check("q5")
```

here's a visualization

```python
# Plot the bar chart
plt.figure(figsize=(12, 6))
monthly_returns.plot(kind='bar')
```

since we can add log returns, we add up all the log returns in the above to see how the strategy does over time

### where do you expect the strategy to do well? when do you expect it to do worse?

```python
monthly_returns.cumsum().plot()
```

### where did  the strategy do well? where did it do worse?

```python

```

### what's missing from this that makes it not a 'backtest'? why wouldn't we make this much money if we started trading this today?

```python

```

# Risk Premia Effect?

maybe the returns at the end of the month are higher because we see bigger movements at the end of the month

how can we test this?

```python
mean_log = tltdf.groupby('dayofmonth').mean()['logreturn']
abs_mean_log = np.abs(mean_log)
plt.bar(abs_mean_log.index, abs_mean_log)
```

### what does the data show? do you think this effect is a risk premia?

```python

```

# Other Assets?

lets generalize some of our functions from before, so they're easier to use in the future

```python

# lets make the downloading and preprocessing data easy to use

def processData(ticker, start_date = "2002-12-31", end_date = "2023-01-01"):
    # get data
    ticker_data = ...
    ticker_raw = ...
    # add log return column
    ticker_raw['logreturn'] = ...
    return ticker_df
    
spydf = processData(ticker="SPY")
spydf

# we can also make our plotting returns by day of month function easy to use so we can run analysis quickly

def plot_mean_dom_return(df):
    df['dayofmonth'] = 
    dom_mean_logret = ...
    plt.bar(..., ...)
    return dom_mean_logret

plot_mean_dom_return(spydf)
    
def df_months(df):
    df_months = ...
    return df_months

df_months(spydf)

spy_monthly_returns = df_months.apply(calculate_returns, 5)

# refer to how you did these things above, so you can apply it here!
```

```python
grader.check("q6")
```

# Other Ideas?

use this notebook as a playground to explore this data and effect further, if you'd like to!

```python

```

```python

```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 2)
```
>>> assert(type(tltraw)==pd.core.frame.DataFrame)
```
```
>>> assert(np.isclose(tltraw['Adj Close'].sum(), 438655.41609191895))
```


Tests for q2 (points: 5)
```
>>> assert(np.isclose(logreturn(50,100),np.log(100/50)))
```
```
>>> assert(np.isclose(simplereturn(50,100),((100-50)/50)))
```
```
>>> assert(np.isclose(simpletolog(1),np.log(100/50)))
```
```
>>> assert(np.isclose(logtosimple(simpletolog(1)),1))
```
```
>>> assert(np.isclose(tltdf['logreturn'].cumsum().values[-1],0.7763374782808363))
```


Tests for q3 (points: 1)
```
>>> assert(np.isclose(tltdf.groupby('dayofmonth').mean()['logreturn'].sum(), 0.005626961357926597))
```


Tests for q4 (points: 3)
```
>>> assert(sep2022['Adj Close'].sum(), 2164.0055618286133)
```
```
>>> assert(last_5_long_922, -0.03122994920153321)
```
```
>>> assert(first_5_short_922, 0.03253599246315185)
```


Tests for q5 (points: 2)
```
>>> assert((calculate_returns(sep2022, 5) < 0.0014) and (calculate_returns(sep2022, 5) > 0.0011))
```
```
>>> assert((monthly_returns.sum() < 1.78) and (monthly_returns.sum() > 1.76))
```


Tests for q6 (points: 2)
```
>>> assert(np.isclose(spydf['Adj Close'].sum(), 859909.4116096497))
```
```
>>> assert((spy_monthly_returns.sum() > -0.119) and (spy_monthly_returns.sum() < -0.117))
```
